# Complexity of Matrix Operations

We will implement different matrix operations by hand, test the implementation and compare the speed with the optimized numpy library. We will also study the computational complexity of these matrix operations, and estimate the performance of our computer.

## Implementation of Matrix Operations (6 points)

Implement the matrix operations $\mathbf{{x}}^T\mathbf{{x}}$, $\mathbf{{A}}\mathbf{{x}}$, $\mathbf{{A}}\mathbf{{A}}$ and $\mathbf{{x}}^T\mathbf{{A}}\mathbf{{x}}$ in Python without using the numpy package. Verify your implementation. 

<!-- grade_id: x_dot_x, points: 0, grade: false, solution: true -->
```python
def x_dot_x(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_x_dot_x, points: 1, grade: true, solution: false -->
```python
assert( x_dot_x( [1,2] ) == 5)
```

<!-- grade_id: A_dot_x, points: 0, grade: false, solution: true -->
```python
def A_dot_x (A, x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_A_dot_x, points: 1, grade: true, solution: false -->
```python
assert( A_dot_x( [ [1,2], [2,1] ], [0,1] ) == [2,1] )
```

<!-- grade_id: A_dot_A, points: 0, grade: false, solution: true -->
```python
def A_dot_A (A):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_A_dot_A, points: 2, grade: true, solution: false -->
```python
assert( A_dot_A( [ [1,1], [2,0] ] ) == [ [3,1], [2,2] ])
```

<!-- grade_id: x_dot_A_dot_x, points: 0, grade: false, solution: true -->
```python
def x_dot_A_dot_x (A,x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_x_dot_A_dot_x, points: 2, grade: true, solution: false -->
```python
assert( x_dot_A_dot_x( [ [1,0], [0,1] ], [2,1] ) == 5)
```

## Measurement and Visualization of Run-Time (7 points)

For each of the implemented matrix operations, measure the execution
time as a function of $n$, up to execution times on the order of one 
second. Compare the execution times of your implementation with a
dedicated matrix library, e.g. the numpy package within python. Plot
the execution times for all matrix operations and both
implementations.

<!-- grade_id: measure, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Present your results in a clear and understandable form. Make sure all features you refer to in the discussion below can easily be identified.

## Interpretation (7 points)

Base your answers to the following questions on your implementation and measurements above. Explain your reasoning. Refer to the plot(s) and other results where appropriate.

How do the runtimes of the implementation in pure Python and numpy compare? Can you explain the differences?

<!-- grade_id: runtime, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Based on the plot(s) from the previous part, compare the computational complexity of the different matrix operations. Do the results agree with your expectations?

<!-- grade_id: complexity, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

How many floating point operations per second do the algorithms achieve? It is sufficient to quote a few examples.

On which hardware did you execute the tests? Are your results in line with the FLOPS of your computer?

<!-- grade_id: cell-d44a100733e0da26, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE