# Frequency Response of an LRC Circuit

Simulate the behavior of a driven LRC circuit with $R$=23$\Omega$, $C$=654nF and $L$=3.2mH.

The LRC circuit can be treated as a voltage divider with complex impedances $Z_R = R$, $Z_C = 1 / (i \omega C)$ and $Z_L=i\omega{}L$. The voltage across one of the components $R,C,L$ is then given by:
$$
  U_{R,C,L} = U_0 \frac{Z_{R,C,L}}{Z_R + Z_C + Z_L}
$$
where $U_0$ is the output voltage of the voltage source. Because $Z_C$ and $Z_L$ depend on the frequency, this calculation cannot be done in the time domain, but only in the frequency domain.

```python
import numpy as np
from numpy import sin,cos,log,log10
import matplotlib.pyplot as plt

R = 23     # ohm
C = 654e-9 # farad
L = 3.2e-3 # henry

nsamples = 1024  # number of samples
fsamp = 1000000  # sampling rate
```

## Fourier spectrum (4 points)

Apply a discrete fourier transform to the following time-descrete signals with 1024 samples taken at a sampling rate of 1 MHz:

- sinusiodal signals $U_0(t) = \sin ( 2 \pi f t )$ with frequencies of 90 kHz and 125 kHz
- square wave signals with frequencies of 15625 Hz and 20000 Hz

Plot the Fourier spectrum of these signals.

<!-- grade_id: fft-plots, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Discuss the Fourier spectra of the four signals from the previous question.

<!-- grade_id: cell-45a45dfb86611eca, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Circuit analysis (3 points)

Determine the ratios of the voltages across resistor ($U_R/U_0$), capacitor ($U_C/U_0$) and inductor ($U_L/U_0$) to the input voltage in the frequency domain. Visualize your results. For which frequencies are the voltage amplitudes over resistor, capacitor and inductor maximal?

<!-- grade_id: cell-66ed14571e080192, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f1b0d23c8c652808, points: 0, grade: false, solution: true -->
```python
# SET THESE VARIABLE TO THE FREQUENCIES WHERE THE RESPECTIVE VOLTAGE IS MAXIMAL
frmax = 0
fcmax = 0
flmax = 0

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: frmax, points: 1, grade: true, solution: false -->
```python
assert(frmax>0)
```

<!-- grade_id: flmax-fcmax, points: 1, grade: true, solution: false -->
```python
assert(fcmax>0)
assert(flmax>0)
```

## Signals in the time domain (3 points)

Calculate the voltages across resistor, capacitor and inductor in the time domain, i.e. as a function of time. Visualize your results for several sinusoidal and square waveforms that you find interesting.

Briefly describe the frequency dependence of the voltages in the circuit. Can you spot any unexpected features?

<!-- grade_id: cell-7056344239fc9b23, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1fe871682dd5110a, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE